<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Estimate standard deviation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The following diagram shows the height in cm of {@ds1@} adults in a village.</p>
<p>
[[jsxgraph]] 
// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});
//Create points of Data Set 1
var lx1 = {#dorx1#};
var ly1 = {#dory#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}
[[/jsxgraph]]
</p>

<p>Estimate the standard deviation of heights
[[input:ans1]]cm [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[Use the following as a guide.
An interval that is 2s wide usually covers about the middle two thirds of the crosses (i.e. x̅ ± s).
An interval that is 4s wide usually covers about the middle 95% of the crosses (i.e. x̅ ± 2s).
The range of the data is usually between 5s and 6s (i.e. x̅ ± 2.5s to x̅ ± 3s).
In this case the mean is about {@mean_real@} and you should be able to estimate the standard deviation close to the real value of {@mr1@}.
<p>Click here to learn more on how estimate standard deviation <a href="https://ecampus.idems.international/mod/resource/view.php?id=13393" download="">"Estimating Standard deviation".</a></p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[mean1:rand(101)+100;
aa:rand([-1,0,1]);
sd1:rand(5)+5;
bb:rand([1/3,1/2,1,2,3]);
ds1:rand(41)+60;
dun1:makelist(round(float(quantile_normal(rand(1.0),mean1,sd1))),i,1,ds1);
dorx1:sort(dun1);
m1:float(std(dorx1));
mr1:dispdp(m1,1);
mean_real: round(mean(dorx1));

maxfreq1:0;
k:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;

xmin:dorx1[1]-2;
xmax:dorx1[ds1]+2+(dorx1[ds1]-dorx1[1])/1.5;
ymin:-2;
ymax:if maxfreq1+2<15 then 15 else maxfreq1+2;

j:1;
dory:makelist(0,i,1,ds1);
dory[1]:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory[i]:j) else (j:1, dory[i]:j);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>mean = {@mean1@}
sd = {@sd1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>mr1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done you are within 10% of the actual standard deviation!</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty>0.1</falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.25</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done that is quite close! You are within 25% of the actual standard deviation. With practice you should be able to estimate to within 10% of the standard deviation, why don't you have another go.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Not close enough! You should be able to estimate to within 25% of the standard deviation, why don't you have another go.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>750076029</deployedseed>
    <deployedseed>446160850</deployedseed>
    <deployedseed>1892374070</deployedseed>
    <deployedseed>2038073284</deployedseed>
    <deployedseed>1263306952</deployedseed>
    <deployedseed>1834225363</deployedseed>
    <deployedseed>71100161</deployedseed>
    <deployedseed>150651908</deployedseed>
    <deployedseed>208727463</deployedseed>
    <deployedseed>1872513520</deployedseed>
    <deployedseed>922251234</deployedseed>
    <deployedseed>1521000575</deployedseed>
    <deployedseed>1146896948</deployedseed>
    <deployedseed>1418421722</deployedseed>
    <deployedseed>1852836895</deployedseed>
    <deployedseed>1692796730</deployedseed>
    <deployedseed>755010582</deployedseed>
    <deployedseed>1263532427</deployedseed>
    <deployedseed>801291460</deployedseed>
    <deployedseed>643940405</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mr1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
