<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Place the estimated value of mean and median in stacked dotplot</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The precipitation in {@ds1@} Eastern African cities in August was measured and the results were grouped in intervals of 6mm. The results are displayed in the graph below. Estimate 'by eye' the values of the mean and median precipitation. Give your answer by dragging the two vertical lines corresponding to the mean and median to the estimated position.</p>

<p>
[[jsxgraph  input-ref-ans1="sR1" input-ref-ans2="sR2" input-ref-ans3="sR3" input-ref-ans4="sR4"]] 
JXG.Options.text.useMathJax = true; 

// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});

// Create points for the Data Set
var lx1 = {#dorx1#};
var ly1 = {#dory#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

// Create hidden point for the mean
var qm1 = board.create('point', [5,0], {
  name: "C",
  size: 2,
  visible: false
});

// Line for the mean
var qm11 = board.create('point', [5,5], {
  name: "C",
  size: 2,
  visible: false
});
var m1 =  board.create('line', [qm1, qm11], {
  straightFirst: true,
  straightLast: true,
  strokeWidth: 2,
  strokeColor: '#21a34c'
});
var lblm1 = board.create('text', [function() {return qm1.X()+1},{#ymax-1#},'mean'], { 
  fontSize: 15,
  color: '#21a34c'
});

// Create hidden point for the median
var qm2 = board.create('point', [{#xmax#}-20,0], {
  name: "C",
  size: 2,
  visible: false
});

// Line for the median
var qm21 = board.create('point', [{#xmax#}-20,5], {
  name: "C",
  size: 2,
  visible: false
});
var m2 =  board.create('line', [qm2,qm21], {
  straightFirst: true,
  straightLast: true,
  strokeWidth: 2,
  strokeColor: '#901B77'
});
var lblm2 = board.create('text', [function() {return qm2.X()+1},{#ymax-2#},'median'], { 
  fontSize: 15,
  color: '#901B77'
});

/* Set answers from state of graph */
stack_jxg.bind_point(sR1,qm1);
stack_jxg.bind_point(sR3,qm11);
stack_jxg.bind_point(sR2,qm2);
stack_jxg.bind_point(sR4,qm21);
board.unsuspendUpdate();


[[/jsxgraph]]
</p>

<p>[[feedback:prt3]]</p>

<p style="display:none;">[[input:ans1]] [[validation:ans1]] </p>

<p>[[feedback:prt1]]</p>

<p style="display:none;">[[input:ans2]] [[validation:ans2]]</p>

<p style="display:none;">[[input:ans3]] [[validation:ans3]]</p>

<p style="display:none;">[[input:ans4]] [[validation:ans4]]</p>

<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Parameters for the distribution */
mean1:rand(30)+30;
sd1:(rand(0.3)+0.5)*mean1;
ds1:rand(50)+50;
bb:(sd1^2/mean1);
aa:(mean1^2/sd1^2);

/* Create the data using with a left skew */
dunx:makelist(round(float(quantile_gamma(rand(1.0),aa,bb))),i,1,ds1);
dorx:sort(dunx);

/* Group data on intervals of length 6 */
dorx1:makelist(1,i,1,ds1);
for i:1 thru ds1 do dorx1[i]:6*floor(dorx[i]/6)+3;

/* Calculate the actual mean of the grouped data */
m1:float(mean(dorx1));
mr1:round(m1)

/* Calculate the actual median of the grouped data */
m2:float(median(dorx1));
mr2:round(m2)

/* Finding bounds for the graph */
/* Calculating the maximum frequency to determine the max y-value */
maxfreq1:0;
k:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;
/* Define max y-value. Set it to a minimum of 20 so that points are not squashed together */
ymax:if maxfreq1+2<20 then 20 else maxfreq1+2;
/* Define the min y-value and the max and min x-values */
xmax:dorx1[ds1]+2;/*+(dorx[ds1]-dorx[1])/1.5 can be added to create empty space on the right to make it visible on mobile phones */
xmin:if dorx1[1]-10<0 then dorx1[1]-10 else -10;
ymin:-2;

/* Find the y-coordinate for each point */
j:1;
dory:makelist(0,i,1,ds1);
dory[1]:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory[i]:j) else (j:1, dory[i]:j);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Mean = "rand(30)+30" = {@mean1@}, SD = "(rand(0.3)+0.5)*Mean" = {@sd1@}</p> 
<p>Data Size = "rand(50)+50" = {@ds1@}</p>
<p>Gamma Parameters [a,b] = [{@aa@},{@bb@}]</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>[m1, 1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>[m2, 1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>[m1, 2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>[m2, 2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1[1]</sans>
        <tans>m1</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the mean is close enough. The actual value is {@mr1@} mm.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1[1]</sans>
        <tans>m1</tans>
        <testoptions>0.2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the mean is almost close enough. The actual value is {@mr1@} mm, with practice you should be able to estimate the mean to within 10%.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>You are a bit too far from the actual mean value of {@mr1@} mm. Your estimate of the mean is more than 20% away, with practice you should be able to do better.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans2[1]</sans>
        <tans>m2</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the median is close enough. The actual value is {@mr2@} mm.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans2[1]</sans>
        <tans>m2</tans>
        <testoptions>0.2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the medain is almost close enough. The actual value is {@mr2@} mm, with practice you should be able to estimate the median to within 10%.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>You are a bit too far from the actual median value of {@mr2@} mm. Your estimate of the median is more than 20% away, with practice you should be able to do better.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[correct_rel_pos:if m1>m2 then (if ans1[1]>ans2[1] then 1 else 0) else (if ans1[1]<ans2[1] then 1 else 0);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>correct_rel_pos</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done you have correctly reflected the skewness in the data by your relative positioning of the mean and median. A quick reminder is that outliers affect the mean but not the median.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your relative positioning of the mean and median do not reflect the skewness in the data. Remember that outliers affect the mean but not the median, hence with skew data the mean is 'pulled' out towards the extreme values while the median is not affected by individual extreme values.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1180180573</deployedseed>
    <deployedseed>897215019</deployedseed>
    <deployedseed>1947941280</deployedseed>
    <deployedseed>231250545</deployedseed>
    <deployedseed>1127514141</deployedseed>
    <deployedseed>1588005898</deployedseed>
    <deployedseed>254115884</deployedseed>
    <deployedseed>46394728</deployedseed>
    <deployedseed>1922285942</deployedseed>
    <deployedseed>1774773390</deployedseed>
    <deployedseed>135245803</deployedseed>
    <deployedseed>1762319639</deployedseed>
    <deployedseed>258514681</deployedseed>
    <deployedseed>133230507</deployedseed>
    <deployedseed>76854659</deployedseed>
    <deployedseed>1705835802</deployedseed>
    <deployedseed>544875161</deployedseed>
    <deployedseed>21686271</deployedseed>
    <deployedseed>1529701231</deployedseed>
    <deployedseed>1971182280</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>[m1, 1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>[m2, 1]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>[m1, 2]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>[m2, 2]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>Stacked Dot Plot</text>
</tag>
      <tag><text>Mean</text>
</tag>
      <tag><text>Median</text>
</tag>
      <tag><text>Estimation</text>
</tag>
    </tags>
  </question>
</quiz>
