<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interpret and match dotplots to respective boxplots</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Below are stacked Dot Plots and Box Plots for four different sets of data.</p> 
<p>Which box plot matches with each dot plot?</p>

<p>Dot plots:</p>
<p>[[jsxgraph]] 
// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});

//Create points of Data Set 4
var lx4 = {#datx4#};
var ly4 = {#daty4#};
var i = 0;
for (i = 0; i <= lx4.length - 1; i++) {
  p[i] = board.create('point', [lx4[i], ly4[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 4
var txt = board.create('text', [{#xmin#}+1, {#maxfreq4+ver_diff-3#}, "D4"], {
  fontSize: 20,fixed: true
});

//Create divisor between V3 and V4
var div1 = board.create('line', [
[{#xmin#},{#maxfreq4+ver_diff+0.5#}],
[{#xmax#},{#maxfreq4+ver_diff+0.5#}]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 1,
dash:2
});

//Create divisor between V2 and V3
var div2 = board.create('line', [
[{#xmin#},{#maxfreq4+maxfreq3+2*ver_diff+0.5#}],
[{#xmax#},{#maxfreq4+maxfreq3+2*ver_diff+0.5#}]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 1,
dash:2
});

//Create divisor between V1 and V2
var div3 = board.create('line', [
[{#xmin#},{#maxfreq4+maxfreq3+maxfreq2+3*ver_diff+0.5#}],
[{#xmax#},{#maxfreq4+maxfreq3+maxfreq2+3*ver_diff+0.5#}]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 1,
dash:2
});

//Create points of Data Set 3
var lx3 = {#datx3#};
var ly3 = {#daty3#};
var i = 0;
for (i = 0; i <= lx3.length - 1; i++) {
  p[i] = board.create('point', [lx3[i], ly3[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 3
var txt = board.create('text', [{#xmin#}+1, {#maxfreq3+maxfreq4+2*ver_diff-3#}, "D3"], {
  fontSize: 20,fixed: true
});

//Create points of Data Set 2
var lx2 = {#datx2#};
var ly2 = {#daty2#};
var i = 0;
for (i = 0; i <= lx2.length - 1; i++) {
  p[i] = board.create('point', [lx2[i], ly2[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 2
var txt = board.create('text', [{#xmin#}+1, {#maxfreq2+maxfreq3+maxfreq4+3*ver_diff-3#}, "D2"], {
  fontSize: 20,fixed: true
});

//Create points of Data Set 1
var lx1 = {#datx1#};
var ly1 = {#daty1#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 1
var txt = board.create('text', [{#xmin#}+1, {#maxfreq1+maxfreq2+maxfreq3+maxfreq4+4*ver_diff-3#}, "D1"], {
  fontSize: 20,fixed: true
});
[[/jsxgraph]]</p>

<p>Box Plots:</p>
<p>
[[jsxgraph]]

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, 10, {#xmax#}, -1],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});

//Create variables for graph 1
var min1 ={#info1[1]#};
var lq1 = {#info1[2]#};
var med1 = {#info1[3]#};
var uq1 = {#info1[4]#};
var max1 = {#info1[5]#};

//Create graph 1
var d11 = board.create('line', [[min1, 7],[min1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d12 = board.create('line', [[max1, 7],[max1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d13 = board.create('line', [[med1, 7],[med1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d14 = board.create('line', [[lq1, 7],[lq1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d15 = board.create('line', [[uq1, 7],[uq1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d16 = board.create('line', [[uq1, 7],[lq1, 7]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d17 = board.create('line', [[uq1, 8],[lq1, 8]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d18 = board.create('line', [[min1, 7.5],[lq1, 7.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d19 = board.create('line', [[max1, 7.5],[uq1, 7.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});

//Create Label for graph 1
var txt = board.create('text', [{#xmin#}+0.5, 8, "A"], {
  fontSize: 20,fixed: true
});

//Create variables for graph 2
var min2 ={#info2[1]#};
var lq2 = {#info2[2]#};
var med2 = {#info2[3]#};
var uq2 = {#info2[4]#};
var max2 = {#info2[5]#};

//Create graph 2
var d21 = board.create('line', [[min2, 5],[min2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d22 = board.create('line', [[max2, 5],[max2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d23 = board.create('line', [[med2, 5],[med2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d24 = board.create('line', [[lq2, 5],[lq2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d25 = board.create('line', [[uq2, 5],[uq2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d26 = board.create('line', [[uq2, 5],[lq2, 5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d27 = board.create('line', [[uq2, 6],[lq2, 6]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d28 = board.create('line', [[min2, 5.5],[lq2, 5.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d29 = board.create('line', [[max2, 5.5],[uq2, 5.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});

//Create Label for graph 2
var txt = board.create('text', [{#xmin#}+0.5, 6, "B"], {
  fontSize: 20,fixed: true
});

//Create variables for graph 3
var min3 ={#info3[1]#};
var lq3 = {#info3[2]#};
var med3 = {#info3[3]#};
var uq3 = {#info3[4]#};
var max3 = {#info3[5]#};

//Create graph 3
var d31 = board.create('line', [[min3, 3],[min3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d32 = board.create('line', [[max3, 3],[max3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d33 = board.create('line', [[med3, 3],[med3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d34 = board.create('line', [[lq3, 3],[lq3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d35 = board.create('line', [[uq3, 3],[uq3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d36 = board.create('line', [[uq3, 3],[lq3, 3]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d37 = board.create('line', [[uq3, 4],[lq3, 4]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d38 = board.create('line', [[min3, 3.5],[lq3, 3.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d39 = board.create('line', [[max3, 3.5],[uq3, 3.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});

//Create Label for graph 3
var txt = board.create('text', [{#xmin#}+0.5, 4, "C"], {
  fontSize: 20,fixed: true
});

//Create variables for graph 4
var min4 ={#info4[1]#};
var lq4 = {#info4[2]#};
var med4 = {#info4[3]#};
var uq4 = {#info4[4]#};
var max4 = {#info4[5]#};

//Create graph 4
var d41 = board.create('line', [[min4, 1],[min4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d42 = board.create('line', [[max4, 1],[max4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d43 = board.create('line', [[med4, 1],[med4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d44 = board.create('line', [[lq4, 1],[lq4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d45 = board.create('line', [[uq4, 1],[uq4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d46 = board.create('line', [[uq4, 1],[lq4, 1]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d47 = board.create('line', [[uq4, 2],[lq4, 2]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d48 = board.create('line', [[min4, 1.5],[lq4, 1.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});
var d49 = board.create('line', [[max4, 1.5],[uq4, 1.5]], {straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true});

//Create Label for graph 4
var txt = board.create('text', [{#xmin#}+0.5, 2, "D"], {
  fontSize: 20,fixed: true
});


[[/jsxgraph]]
</p>

<p>Enter the relevant letter for each Stacked Dot Plot:</p>

<p>\(D1\) matches with: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>\(D2\) matches with: [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>\(D3\) matches with: [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
<p>\(D4\) matches with: [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Remember how a boxplot is defined.</p>
<p>The middle of the box is the median line. Half the data are above the median and half are below. Looking at the dotplots we see that some of the distributions are skewed, so the median point is very low or very high. Others are more symmetric, so the boxplot will also be more symmetric.</p>
<p>Where are the 25% (lower quartile) and 75% (upper quartile) points? Then match these to the boxplots.</p>
<p>One of the datasets looks "bimodal" meaning it has two "peaks". How does a boxplot of bimodal data look different from a boxplot of a normally distributed data set?</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[dat_min: rand(101)+100;
dat_max: dat_min+30+rand(21);
n: rand(41)+60;
bin_width: floor((dat_max-dat_min)/30);
ver_diff: 2;

/* dat1: normal */
dat1: makelist(float(quantile_normal(rand(1.0),0,1)),i,1,n);
dat1: ((dat1-lmin(dat1))/(lmax(dat1)-lmin(dat1)))*(dat_max-dat_min)+dat_min;
dat1: sort(dat1);
dat1: round(dat1);
datx1:makelist(1,i,1,n);
for i:1 thru n do datx1[i]:bin_width*floor(dat1[i]/bin_width)+bin_width/2;

/*  info of normal for box plot: [min,lq,med,uq,max] */
min1:datx1[1];
lq1:quantile (datx1,1/4),numer;
med1:median(datx1);
uq1:quantile (datx1, 3/4),numer;
max1:datx1[n];
info1:[min1,lq1,med1,uq1,max1];

/* dat2: bimodal (from two disjoint normals) */
dat21: makelist(float(quantile_normal(rand(1.0),0,1)),i,1,floor(n/2));
dat22: makelist(float(quantile_normal(rand(1.0),4,1)),i,1,ceiling(n/2));
dat2: append(dat21,dat22);
dat2: ((dat2-lmin(dat2))/(lmax(dat2)-lmin(dat2)))*(dat_max-dat_min)+dat_min;
dat2: sort(dat2);
dat2: round(dat2);
datx2:makelist(1,i,1,n);
for i:1 thru n do datx2[i]:bin_width*floor(dat2[i]/bin_width)+bin_width/2;

/*  info of bimodal for box plot: [min,lq,med,uq,max] */
min2:datx2[1];
lq2:quantile (datx2,1/4),numer;
med2:median(datx2);
uq2:quantile (datx2, 3/4),numer;
max2:datx2[n];
info2:[min2,lq2,med2,uq2,max2];

/* dat3: gamma right skewed */
dat3: makelist(float(quantile_gamma(rand(1.0),2,20)),i,1,n);
dat3: ((dat3-lmin(dat3))/(lmax(dat3)-lmin(dat3)))*(dat_max-dat_min)+dat_min;
dat3: sort(dat3);
dat3: round(dat3);
datx3:makelist(1,i,1,n);
for i:1 thru n do datx3[i]:bin_width*floor(dat3[i]/bin_width)+bin_width/2;

/*  info of right skew for box plot: [min,lq,med,uq,max] */
min3:datx3[1];
lq3:quantile (datx3,1/4),numer;
med3:median(datx3);
uq3:quantile (datx3, 3/4),numer;
max3:datx3[n];
info3:[min3,lq3,med3,uq3,max3];

/* dat4: gamma left skewed */
dat4: makelist(float(quantile_gamma(rand(1.0),2,20)),i,1,n);
dat4: lmax(dat4)-lmin(dat4)+2-dat4;
dat4: ((dat4-lmin(dat4))/(lmax(dat4)-lmin(dat4)))*(dat_max-dat_min)+dat_min;
dat4: sort(dat4);
dat4: round(dat4);
datx4:makelist(1,i,1,n);
for i:1 thru n do datx4[i]:bin_width*floor(dat4[i]/bin_width)+bin_width/2;

/*  info of left skew for box plot: [min,lq,med,uq,max] */
min4:datx4[1];
lq4:quantile (datx4,1/4),numer;
med4:median(datx4);
uq4:quantile (datx4, 3/4),numer;
max4:datx4[n];
info4:[min4,lq4,med4,uq4,max4];

/* Randomising the order of the graphs */
dumlist1:[1,2,3,4];
randdumlist1:random_permutation(dumlist1);
datasetslist:[datx1,datx2,datx3,datx4];
datx1:datasetslist[randdumlist1[1]];
datx2:datasetslist[randdumlist1[2]];
datx3:datasetslist[randdumlist1[3]];
datx4:datasetslist[randdumlist1[4]];

/* version number for each graph */
normalans1:if randdumlist1[1] = 1 then 1 elseif randdumlist1[2]=1 then 2 elseif randdumlist1[3]=1 then 3 else 4;
bimodalans1:if randdumlist1[1] = 2 then 1 elseif randdumlist1[2]=2 then 2 elseif randdumlist1[3]=2 then 3 else 4;
rightskewans1:if randdumlist1[1] = 3 then 1 elseif randdumlist1[2]=3 then 2 elseif randdumlist1[3]=3 then 3 else 4;
leftskewans1:if randdumlist1[1] = 4 then 1 elseif randdumlist1[2]=4 then 2 elseif randdumlist1[3]=4 then 3 else 4;

/* Randomising the order of the graphs */
dumlist2:[1,2,3,4];
randdumlist2:random_permutation(dumlist2);
infolist:[info1,info2,info3,info4];
info1:infolist[randdumlist2[1]];
info2:infolist[randdumlist2[2]];
info3:infolist[randdumlist2[3]];
info4:infolist[randdumlist2[4]];

/* version number for each graph */
normalans2:if randdumlist2[1] = 1 then 1 elseif randdumlist2[2]=1 then 2 elseif randdumlist2[3]=1 then 3 else 4;
bimodalans2:if randdumlist2[1] = 2 then 1 elseif randdumlist2[2]=2 then 2 elseif randdumlist2[3]=2 then 3 else 4;
rightskewans2:if randdumlist2[1] = 3 then 1 elseif randdumlist2[2]=3 then 2 elseif randdumlist2[3]=3 then 3 else 4;
leftskewans2:if randdumlist2[1] = 4 then 1 elseif randdumlist2[2]=4 then 2 elseif randdumlist2[3]=4 then 3 else 4;

/* Calculate the maximun frequency of each data set */
maxfreq1:0;
k:1;
for i:2 thru n do if datx1[i]=datx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;

maxfreq2:0;
k:1;
for i:2 thru n do if datx2[i]=datx2[i-1] then (k:k+1, if maxfreq2 < k then maxfreq2:k) else k:1;

maxfreq3:0;
k:1;
for i:2 thru n do if datx3[i]=datx3[i-1] then (k:k+1, if maxfreq3 < k then maxfreq3:k) else k:1;

maxfreq4:0;
k:1;
for i:2 thru n do if datx4[i]=datx4[i-1] then (k:k+1, if maxfreq4 < k then maxfreq4:k) else k:1;

/* Begin generating stacked dot plots */

xmin:dat_min-2;
xmax:dat_max+2;
ymin:-4;
ymax: maxfreq1+maxfreq2+maxfreq3+maxfreq4+3*ver_diff+1;

/* Generate y-values for each data set to plot the stacked dots */
j:1;
daty1:makelist(0,i,1,n);
daty1[1]:1+maxfreq2+maxfreq3+maxfreq4+3*ver_diff;
for i:2 thru n do if datx1[i]=datx1[i-1] then (j:j+1, daty1[i]:j+maxfreq2+maxfreq3+maxfreq4+3*ver_diff) else (j:1, daty1[i]:j+maxfreq2+maxfreq3+maxfreq4+3*ver_diff);

j:1;
daty2:makelist(0,i,1,n);
daty2[1]:1+maxfreq3+maxfreq4+2*ver_diff;
for i:2 thru n do if datx2[i]=datx2[i-1] then (j:j+1, daty2[i]:j+maxfreq3+maxfreq4+2*ver_diff) else (j:1, daty2[i]:j+maxfreq3+maxfreq4+2*ver_diff);

j:1;
daty3:makelist(0,i,1,n);
daty3[1]:1+maxfreq4+ver_diff;
for i:2 thru n do if datx3[i]=datx3[i-1] then (j:j+1, daty3[i]:j+maxfreq4+ver_diff) else (j:1, daty3[i]:j+maxfreq4+ver_diff);

j:1;
daty4:makelist(0,i,1,n);
daty4[1]:1;
for i:2 thru n do if datx4[i]=datx4[i-1] then (j:j+1, daty4[i]:j) else (j:1, daty4[i]:j);

/* End generating stacked dot plots */
ta1_c:if randdumlist1[1]=randdumlist2[1] then "A" elseif randdumlist1[1]=randdumlist2[2] then "B" elseif randdumlist1[1]=randdumlist2[3] then "C" else "D";
ta1:[["A",is(ta1_c="A"), "Boxplot A"],["B",is(ta1_c="B"), "Boxplot B"],["C",is(ta1_c="C"), "Boxplot C"],["D",is(ta1_c="D"), "Boxplot D"]];
ta2_c:if randdumlist1[2]=randdumlist2[1] then "A" elseif randdumlist1[2]=randdumlist2[2] then "B" elseif randdumlist1[2]=randdumlist2[3] then "C" else "D";
ta2:[["A",is(ta2_c="A"), "Boxplot A"],["B",is(ta2_c="B"), "Boxplot B"],["C",is(ta2_c="C"), "Boxplot C"],["D",is(ta2_c="D"), "Boxplot D"]];
ta3_c:if randdumlist1[3]=randdumlist2[1] then "A" elseif randdumlist1[3]=randdumlist2[2] then "B" elseif randdumlist1[3]=randdumlist2[3] then "C" else "D";
ta3:[["A",is(ta3_c="A"), "Boxplot A"],["B",is(ta3_c="B"), "Boxplot B"],["C",is(ta3_c="C"), "Boxplot C"],["D",is(ta3_c="D"), "Boxplot D"]];
ta4_c:if randdumlist1[4]=randdumlist2[1] then "A" elseif randdumlist1[4]=randdumlist2[2] then "B" elseif randdumlist1[4]=randdumlist2[3] then "C" else "D";
ta4:[["A",is(ta4_c="A"), "Boxplot A"],["B",is(ta4_c="B"), "Boxplot B"],["C",is(ta4_c="C"), "Boxplot C"],["D",is(ta4_c="D"), "Boxplot D"]]]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Data size = {@n@}. Answer list: {@[ta1_c,ta2_c,ta3_c,ta4_c]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1801006757</deployedseed>
    <deployedseed>768252992</deployedseed>
    <deployedseed>1901898753</deployedseed>
    <deployedseed>1925301602</deployedseed>
    <deployedseed>1316343442</deployedseed>
    <deployedseed>2109495421</deployedseed>
    <deployedseed>1543454398</deployedseed>
    <deployedseed>140518220</deployedseed>
    <deployedseed>1665582679</deployedseed>
    <deployedseed>1030164503</deployedseed>
    <deployedseed>425942596</deployedseed>
    <deployedseed>451148150</deployedseed>
    <deployedseed>1369728296</deployedseed>
    <deployedseed>374485637</deployedseed>
    <deployedseed>1776053810</deployedseed>
    <deployedseed>847030446</deployedseed>
    <deployedseed>271459631</deployedseed>
    <deployedseed>1522679322</deployedseed>
    <deployedseed>1388158991</deployedseed>
    <deployedseed>57289455</deployedseed>
    <deployedseed>553761874</deployedseed>
    <deployedseed>2006693680</deployedseed>
    <deployedseed>335428334</deployedseed>
    <deployedseed>1807555479</deployedseed>
    <deployedseed>336215593</deployedseed>
    <deployedseed>1038662071</deployedseed>
    <deployedseed>2050421240</deployedseed>
    <deployedseed>854347360</deployedseed>
    <deployedseed>616343465</deployedseed>
    <deployedseed>1496310453</deployedseed>
    <deployedseed>819154543</deployedseed>
    <deployedseed>493731052</deployedseed>
    <deployedseed>1451165897</deployedseed>
    <deployedseed>1036523842</deployedseed>
    <deployedseed>471844676</deployedseed>
    <deployedseed>528939971</deployedseed>
    <deployedseed>3402304</deployedseed>
    <deployedseed>1106143636</deployedseed>
    <deployedseed>1576038739</deployedseed>
    <deployedseed>1619350419</deployedseed>
    <deployedseed>1149313236</deployedseed>
    <deployedseed>300087478</deployedseed>
    <deployedseed>1279523893</deployedseed>
    <deployedseed>1539829830</deployedseed>
    <deployedseed>830626927</deployedseed>
    <deployedseed>1041797952</deployedseed>
    <deployedseed>740558512</deployedseed>
    <deployedseed>1062067506</deployedseed>
    <deployedseed>1801148755</deployedseed>
    <deployedseed>1403585311</deployedseed>
    <deployedseed>1909028018</deployedseed>
    <deployedseed>1235618568</deployedseed>
    <deployedseed>1300948314</deployedseed>
    <deployedseed>1667452766</deployedseed>
    <deployedseed>173301734</deployedseed>
    <deployedseed>56555878</deployedseed>
    <deployedseed>440003830</deployedseed>
    <deployedseed>576180509</deployedseed>
    <deployedseed>1952330545</deployedseed>
    <deployedseed>219751550</deployedseed>
    <deployedseed>1478334696</deployedseed>
    <deployedseed>81902505</deployedseed>
    <deployedseed>102107907</deployedseed>
    <deployedseed>1141627915</deployedseed>
    <deployedseed>1728031519</deployedseed>
    <deployedseed>769958447</deployedseed>
    <deployedseed>1657940095</deployedseed>
    <deployedseed>483317018</deployedseed>
    <deployedseed>1327974645</deployedseed>
    <deployedseed>1588417208</deployedseed>
    <deployedseed>1152955789</deployedseed>
    <deployedseed>55023231</deployedseed>
    <deployedseed>1319737195</deployedseed>
    <deployedseed>1012899027</deployedseed>
    <deployedseed>678812362</deployedseed>
    <deployedseed>1651285347</deployedseed>
    <deployedseed>724032048</deployedseed>
    <deployedseed>1334280520</deployedseed>
    <deployedseed>735003663</deployedseed>
    <deployedseed>1445763220</deployedseed>
    <deployedseed>1556214072</deployedseed>
    <deployedseed>1041586978</deployedseed>
    <deployedseed>466667189</deployedseed>
    <deployedseed>1559231113</deployedseed>
    <deployedseed>2116649457</deployedseed>
    <deployedseed>931173427</deployedseed>
    <deployedseed>1401011273</deployedseed>
    <deployedseed>418278936</deployedseed>
    <deployedseed>999781913</deployedseed>
    <deployedseed>350490041</deployedseed>
    <deployedseed>1270066446</deployedseed>
    <deployedseed>1011255394</deployedseed>
    <deployedseed>1970947858</deployedseed>
    <deployedseed>1959094721</deployedseed>
    <deployedseed>1799029378</deployedseed>
    <deployedseed>400906966</deployedseed>
    <deployedseed>1224324728</deployedseed>
    <deployedseed>1647864254</deployedseed>
    <deployedseed>1271249744</deployedseed>
    <deployedseed>1889491786</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
